---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Parallel Processing for Larger Datasets

```yaml
type: VideoExercise
key: ceae36edcb
xp: 50
```

`@projector_key`
25e3bda2b8a26bbfb2f6c8fdafeb3eca
